# Week1 Assignment - Testing OKPy and the autograder.

For this assignment, you will load the iris flowers dataset from sklearn and predict flower species based on their petal and sepal sizes.

```python
"""This is a docstring!"""
```

$\beta=5$

```python
alpha = 4
```

```python
import pathlib
p = pathlib.Path()
```

```python
a = p.absolute() / ""
```

Output:
```
pathlib.PosixPath
```

```python
var_list = %who_ls
```

```python
var_list
```

Output:
```
['datasets', 'iris']
```

## Part A. Load the Iris dataset.

Sklearn comes with prepared datasets that you can load. To load the iris dataset, run:
```python
from sklearn import datasets
iris = datasets.load_iris()
```

```python
from sklearn import datasets
iris = datasets.load_iris()
```

```python
%lsmagic #list all the magic commands 
```

Output:
```
Available line magics:
%alias  %alias_magic  %autoawait  %autocall  %automagic  %autosave  %bookmark  %cat  %cd  %clear  %colors  %conda  %config  %connect_info  %cp  %debug  %dhist  %dirs  %doctest_mode  %ed  %edit  %env  %gui  %hist  %history  %killbgscripts  %ldir  %less  %lf  %lk  %ll  %load  %load_ext  %loadpy  %logoff  %logon  %logstart  %logstate  %logstop  %ls  %lsmagic  %lx  %macro  %magic  %man  %matplotlib  %mkdir  %more  %mv  %notebook  %page  %pastebin  %pdb  %pdef  %pdoc  %pfile  %pinfo  %pinfo2  %pip  %popd  %pprint  %precision  %prun  %psearch  %psource  %pushd  %pwd  %pycat  %pylab  %qtconsole  %quickref  %recall  %rehashx  %reload_ext  %rep  %rerun  %reset  %reset_selective  %rm  %rmdir  %run  %save  %sc  %set_env  %store  %sx  %system  %tb  %time  %timeit  %unalias  %unload_ext  %who  %who_ls  %whos  %xdel  %xmode

Available cell magics:
%%!  %%HTML  %%SVG  %%bash  %%capture  %%debug  %%file  %%html  %%javascript  %%js  %%latex  %%markdown  %%perl  %%prun  %%pypy  %%
```

```python
from pathlib import Path
import glob
a = glob.glob('/Users/Alex/Desktop/*')
[j.split('/')[-1] for j in a]
# [j.read_text().split]

```

Output:
```
['Screen Shot 2020-02-02 at 4.12.15 PM.png',
 'graduation',
 'Screen Shot 2020-02-02 at 3.51.58 PM.png',
 'cleanup_pics',
 'Screen Shot 2020-01-29 at 3.59.21 PM.png',
 'tax_9:2018.pdf',
 'Screen Shot 2020-02-02 at 3.50.40 PM.png',
 'Screen Shot 2020-02-12 at 5.58.21 PM.png',
 'Screen Shot 2020-01-26 at 12.24.25 PM.png',
 'Screen Shot 2020-02-02 at 4.14.17 PM.png',
 'test1_results.pdf',
 'Screen Shot 2020-01-26 at 12.11.35 PM.png',
 'Screen Shot 2020-02-11 at 3.06.20 PM.png',
 'Screen Shot 2020-02-02 at 3.52.51 PM.png',
 'Screen Shot 2020-02-01 at 12.40.21 PM.png',
 'Screen Shot 2020-02-02 at 3.53.22 PM.png',
 'mmep',
 'Screen Shot 2020-02-02 at 4.13.19 PM.png',
 'Screen Shot 2020-02-09 at 8.45.34 PM.png',
 'fav',
 'Screen Shot 2020-02-09 at 5.49.12 PM.png',
 'Screen Shot 2020-02-02 at 4.02.34 PM.png',
 'Sofrito Recipe.docx',
 'Screen Shot 2020-02-02 at 3.27.27 PM.png',
 'Screen Shot 2020-02-01 at 12.39.43 PM.png',
 '3d_plot',
 'Screen Shot 2020-01-22 at 3.11.14 PM.png',
 'Screen Shot 2
```

```python
x = !echo hello world
# ! this is calling for the system code (in our case, bash)
x
```

Output:
```
['hello world']
```

```python
import pandas as pd 
df = pd.DataFrame(iris.data, columns=iris.feature_names)
```

```python
df.columns = [s.replace(" ","_").replace("(","").replace(")","") for s in df.columns]
df.head(2)
```

Output:
```
sepal_length_cm  sepal_width_cm  petal_length_cm  petal_width_cm
0              5.1             3.5              1.4             0.2
1              4.9             3.0              1.4             0.2
```

<!-- grade_id: cell-c1e7f79ad332b8f9, points: 0, grade: false, solution: true -->
```python
from sklearn import datasets

# Load the iris dataset
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-37-e00ec749e977> in <module>
      6 iris = datasets.load_iris()
      7 
----> 8 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-fd50146eeead0c65, points: 1, grade: true, solution: false -->
```python
assert (iris['data'] == datasets.load_iris()['data']).all()
```

## Part B. Use an SVM to predict iris species.

To setup the data and train an SVM, run the following:
```python
from sklearn.svm import SVC

X = iris.data
y = iris.target

classifier = SVC(kernel="rbf")
classifier.fit(X, y)
predictions = classifier.predict(X)
```

```python
%who
```

Output:
```
SVC	 X	 classifier	 datasets	 iris	 predictions	 var_list	 y
```

```python
from sklearn.svm import SVC
```

```python
svc = SVC()  #little svs is now is instance of the SVC class. 
svc
```

Output:
```
SVC(C=1.0, break_ties=False, cache_size=200, class_weight=None, coef0=0.0,
    decision_function_shape='ovr', degree=3, gamma='scale', kernel='rbf',
    max_iter=-1, probability=False, random_state=None, shrinking=True,
    tol=0.001, verbose=False)
```

```python
X = iris.data
y = iris.target
```

```python
classifier = SVC()
```

```python
X.shape
```

Output:
```
(150, 4)
```

```python
classifier.fit(X, y)
```

Output:
```
SVC(C=1.0, break_ties=False, cache_size=200, class_weight=None, coef0=0.0,
    decision_function_shape='ovr', degree=3, gamma='scale', kernel='rbf',
    max_iter=-1, probability=False, random_state=None, shrinking=True,
    tol=0.001, verbose=False)
```

```python
classifier.score(X,y) 
```

Output:
```
0.9733333333333334
```

Classification : the outcome is discrete grouping 
regression : the outcome is continuous

```python
predictions = classifier.predict(X)
predictions

```

Output:
```
array([0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0,
       0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0,
       0, 0, 0, 0, 0, 0, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1,
       1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 2, 1, 1, 1, 1, 1, 2, 1, 1, 1, 1,
       1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 2, 2, 2, 2, 2, 2, 1, 2, 2, 2,
       2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2,
       2, 2, 2, 2, 2, 2, 1, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2])
```

- Overfitting (Variance): tailored to a specific dataset (does not generalize well), misses the point, memorizes the data
- Underfitting (Bias): accuracy could be better, too general 

1. Import
2. Instantiate (a class)
  - Estimators:
    - Classifiers
    - Regressors
3. Fit
4. Predict

```python
predictions
```

Output:
```
array([0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0,
       0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0,
       0, 0, 0, 0, 0, 0, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1,
       1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 2, 1, 1, 1, 1, 1, 2, 1, 1, 1, 1,
       1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2,
       2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2,
       2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2])
```

```python
from sklearn.model_selection import train_test_split
```

<!-- grade_id: cell-bc8ce57a60e5fd7a, points: 0, grade: false, solution: true -->
```python
from sklearn.svm import SVC

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-165c467d2a6e7dde, points: 1, grade: true, solution: false -->
```python
assert X.shape == (150, 4)
assert len(y) == 150
```